---
title: "`r emo::ji('woman_teacher')` building interactive tutorials in R"
subtitle: "<br> `r emo::ji('link')` [bit.ly/teach-r-online-mats](https://bit.ly/teach-r-online-mats)"
author: "<br> <br> dr. mine çetinkaya-rundel <br> dr. colin rundel"
output:
  xaringan::moon_reader:
    lib_dir: libs
    css: ["xaringan-themer.css", "slides-style.css"]
    nature:
      highlightStyle: solarized-light
      highlightLines: true
      countIncrementalSlides: false
      ratio: "16:9"
      slideNumberFormat: |
        <div class="progress-bar-container">
          <div class="progress-bar" style="width: calc(%current% / %total% * 100%);">
          </div>
        </div>
---

```{r child = "../setup.Rmd"}
```

```{r set-theme, include=FALSE}
library(xaringanthemer)
style_duo_accent(
  primary_color      = "#7A8387", # pantone monument
  secondary_color    = "#B6CADA", # pantone baby blue
  header_font_google = google_font("Raleway"),
  text_font_google   = google_font("Raleway", "300", "300i"),
  code_font_google   = google_font("Source Code Pro"),
  text_font_size     = "30px"
)
```


```{r load-packages, message=FALSE, echo=FALSE}
library(tidyverse)
library(broom)
library(scales)
library(openintro)
library(jsonlite)
library(usethis)
```

## Activity 

**TO DO: UPDATE**

.hand[while we wait to get started...]

- Go to [INSERT LINK](INSERT LINK)
- Start working through the tutorial
- Feel free to make mistakes and test out the feedback
- If you get to the very end, follow the instructions (but if they seem a bit opaque, don't fret, we'll say more about "submission" later...)

---

.huge-text[you...]

.large[
- know and teach R
- are familiar with R Markdown
- are interested in providing automated feedback
- might be interested in automated marking
]

---

class: middle, inverse

.pull-left[
  .huge-text[why]
]
.pull-right[
  .larger[
  auto  
  feedback
  ]
]

---

class: middle

.pull-left[
.center[
.huge[`r emo::ji("point_right")`] <br>
.large[**Nudging**] <br>
students towards the right answer, especially in formative assessments
]
]

---

```{r include=FALSE}
set.seed(1234)
means <- rnorm(10) %>% round(2)
```

.sample-question[
Suppose 10 means from a simulated sampling distribution is stored in a vector called `means`.
```{r}
means
```
What is the value of the first mean?
]
--
<br>
.pull-left-wide[
.sample-answer[
```{r error=TRUE}
mean[1]
```
]
]
--
.pull-right-narrow[
```{r echo=FALSE, fig.align="right"}
knitr::include_graphics("images/student.png")
```
]

---

## Nudging

.hand[Not all feedback is useful, at least not for beginners...]

<br>

Providing helpful feedback can help them nudge them towards success:

```{r eval=FALSE}
mean[1]
```

```{r echo=FALSE}
ui_oops("`mean` is a function and a function doesn't have elements that can be subsetted with square brackets.")
ui_info("`means` is the vector of sample means calculated earlier.")
```

---

.sample-question[
Visualise the relationship between city and highway mileage of cars from the `mpg` dataset, conditional on year of manufacture.
]

--
<br>
.sample-answer[
There is a strong, positive, linear relationship between the city and highway mileage of cars. Year does not seem to be related to either variable.
```{r message=FALSE, fig.height=2.5, fig.width=9, dpi=300, out.width="70%"}
ggplot(mpg,aes(x = hwy, y = cty, fill=year)) +geom_point()+geom_smooth()
```
]

---

class: middle

.pull-left-wide[
.sample-feedback[
- You mention a linear relationship, however your plot uses a loess fit to visualise the relationship between city and highway mileage. Also, the plot displays the uncertainty around the fit, but you haven't addressed it in your narrative.
- Year should be mapped to the `color` aesthetic, not `fill`. 
- Plot styling: Use informative axis labels, noting units of measurement. Also, give an informative title to your plot.
- Code styling: Use consistent spacing around operators (e.g ` = `) and line breaks after `+` in each layer of your ggplot.
]
]
.pull-right-narrow[
```{r echo=FALSE, fig.align="center"}
knitr::include_graphics("images/teacher.png")
```
]

---

class: middle

.pull-left[
.center[
.huge[`r emo::ji("point_right")`] <br>
.large[**Nudging**] <br>
students towards the right answer, especially in formative assessments
]
]
.pull-right[
.center[
.huge[`r emo::ji("heavy_plus_sign")`] <br>
.large[**Scaling**] <br>
up efficiency of grading faster than (human) resources
]
]

---

## Scaling

.hand[Our courses are growing, and that's a good thing, right?]

--

- Students turning in their work as R Markdown documents makes collecting submissions including code and narrative straightforward.

--

- Providing feedback on both the code and narrative is not scalable unless (human) resources dedicated to your course grow proportionally with enrolments.

---

class: middle, inverse

.pull-left[
  .huge-text[how]
]
.pull-right[
  .larger[
  auto  
  feedback
  ]
]

---

## `r emo::ji("package")` learnr

- **learnr** is an R package that makes it easy to create interactive tutorials from R Markdown documents.

- Tutorials can include:
  - Narrative, figures, illustrations, and equations
  - Code exercises (R code chunks that users can edit and execute directly)
  - Multiple choice questions
  - Videos (YouTube, Vimeo)
  - Interactive Shiny components

- learnr is on CRAN

```{r eval=FALSE}
install.packages("learnr")
```
  
---

background-image: url(images/learnr-first-look.png)
background-position: center
background-repeat: no-repeat
background-size: contain
class: middle

---

## `r emo::ji("package")` gradethis

- Companion to the learnr package, **gradethis** provides multiple methods to grade learnr exercises:
  - `grade_code()`: Grade code against a solution
  - `grade_conditions()`: Grade all specified conditions
  - `grade_result()`: Grade result of exercise code

- gradethis is not yet on CRAN

```{r eval=FALSE}
devtools::install_github("rstudio/gradethis")
```

---

class: middle

.huge[demo]

[**`[tutorial]`**](https://minecr.shinyapps.io/nobel-laureates-learnr) 

[**`[code]`**](https://github.com/mine-cetinkaya-rundel/feed-me-back/blob/master/examples/01-learnr-gradethis/nobel-laureates-learnr.Rmd)

---

class: middle, inverse

.huge-text[sharing]

---

## Sharing with students

- You could share the R Markdown document (and all accompanying files) but that’s probably not what you want to do...

- Deploy on
  - shinyapps.io
  - RStudio Connect (free for academic use, requires setup)

- Road less travelled: distribute as a package

---

class: middle, inverse

.huge-text[recording]
.huge-text[data]

---

.large[
.hand[Q: Is it possible to record student data?]
]

Yes, though not trivial, at this point. 

---

## Recording Attempts

A "good enough" solution for formative exercises: embed a 
Google/Microsoft/etc. Form at the end and ask students to "submit" their work.

This only records that the student reached the end of the tutorial and not how (or even if) they answered any of the questions or exercises.

**Tip:** Adding a free-text question to the form asking students to reflect on the 
exercises they just completed - this is then recorded and can be analyzed. [**`[example]`**](https://minecr.shinyapps.io/04-wrangleviz/#section-finish-up)

---

## Recording Solutions

The [`learnrhash`](https://github.com/rundel/learnrhash) package builds on the previous method by providing a way for students to submit their answers by generating a text "hash" which can be copy and pasted into the web form.

This package is being used to create the submission tool at the end of the sample tutorial.


.footnote[
See also the [**submitter**](https://github.com/dtkaplan/submitr) package by Danny Kaplan.
]

---

.huge-text[demo]

.large[
[**`[tutorial]`**](https://minecr.shinyapps.io/nobel-laureates-learnr) 

[**`[code]`**](https://github.com/mine-cetinkaya-rundel/teach-r-online/blob/master/02-tutorial/examples/lego-sales.Rmd)
]

---

## Submissions

Since I've just submitted, my details and my hash will now be stored in a Google Sheet - if you would like to play around with this data (Instructor mode) you can access it here:

.center[ .large[
https://bit.ly/tutorial-data
] ]

We don't have time today to demo reading in and decoding these data, but we will include code for the whole process in the [learnrhash package repository](https://github.com/rundel/learnrhash) in the next week, using all of the data you submit.



---

class: middle, inverse

.huge-text[closing]
.huge-text[thoughts]

---

## Best practices for automated feedback

--

- Measure twice, cut once (verify the correctness of your tests) 📏📏 ✂️

--

- Use rounding & type coercion to write robust tests 💪

--

- Don't give automated feedback on everything, asking narrative questions that can't be auto checked but gets the student thinking and writing has pedagogical benefits 🙈

--

- Consider peer feedback where automated feedback is not feasible (e.g. interpretation, narrative) but scalability is an issue 👯

---

.large[
.hand[Q: What is an approachable way to get started?]
]

Build a tutorial where students build develop their analysis in exercise code 
chunks (that are not checked) and only multiple choice questions are used for 
assessment. [**`[example]`**](https://minecr.shinyapps.io/04-wrangleviz)

```{r echo=FALSE, out.width="85%", fig.align="center"}
knitr::include_graphics("images/learnr-simple.png")
```

---

.large[
.hand[Q: I've built simple tutorials already. How do I make the jump to code checking and providing automated feedback that is actually useful?]
]

- Replicate `gradethis::gradethis_demo()`, then make incremental changes
- Read the [Testing](https://r-pkgs.org/tests.html#test-tests) chapter in R Packages (Wickham and Bryan) 
- Also read the [Metaprogramming](https://adv-r.hadley.nz/metaprogramming.html) section in Advanced R (Wickham)

```{r echo=FALSE, out.width="75%", fig.align="center"}
knitr::include_graphics("images/learnr-learn-more.jpeg")
```

---

.large[
.hand[Q: Sounds great, but can it handle my class size and usage?]
]

- First, chances are you're not using these live, but you might be...
- If so, make sure to max out your instances and instance size on shinyapps.io.
- Essential reading:
  - [Publishing learnr Tutorials on shinyapps.io](https://cran.r-project.org/web/packages/learnr/vignettes/shinyapps-publishing.html) by Angela Li
  - [Teach R with learnr: a powerful tool for remote teaching](https://education.rstudio.com/blog/2020/05/learnr-for-remote/) by Allison Horst


---

.hand-large[thank you!]

- All materials at [bit.ly/teach-r-online-mats](https://bit.ly/teach-r-online-mats)

- Sign up for the upcoming workshop at [bit.ly/teach-r-online](https://bit.ly/teach-r-online):
  - 17 July: Teaching computing with Git and GitHub 
